---
title       : Data Visualization via ggplot2
description : Insert the chapter description here


--- type:NormalExercise lang:r xp:100 skills:1 key:0b7d52dd7f
## Produce incorrect plot



*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
